Adaptive behavior is guided by integrated representations of controlled and non-controlled information

Author response:

The following is the authors’ response to the original reviews.

eLife Assessment

This useful study uses creative scalp EEG decoding methods to attempt to demonstrate that two forms of learned associations in a Stroop task are dissociable, despite sharing similar temporal dynamics. However, the evidence supporting the conclusions is incomplete due to concerns with the experimental design and methodology. This paper would be of interest to researchers studying cognitive control and adaptive behavior, if the concerns raised in the reviews can be addressed satisfactorily.

We thank the editors and the reviewers for their positive assessment of our work and for providing us with an opportunity to strengthen this manuscript. Please see below our responses to each comment raised in the reviews.

Public Reviews:

Reviewer #1 (Public review):

Summary:

This study focuses on characterizing the EEG correlates of item-specific proportion congruency effects. In particular, two types of learned associations are characterized. One being associations between stimulus features and control states (SC), and the other being stimulus features and responses (SR). Decoding methods are used to identify SC and SR correlates and to determine whether they have similar topographies and dynamics.

The results suggest SC and SR associations are simultaneously coactivated and have shared topographies, with the inference being that these associations may share a common generator.

Strengths:

Fearless, creative use of EEG decoding to test tricky hypotheses regarding latent associations. Nice idea to orthogonalize the ISPC condition (MC/MI) from stimulus features.

Thank you for acknowledging the strength in EEG decoding and design. We have addressed all your concerns raised below point by point.

Weaknesses:

(1a) I'm relatively concerned that these results may be spurious. I hope to be proven wrong, but I would suggest taking another look at a few things.

While a nice idea in principle, the ISPC manipulation seems to be quite confounded with the trial number. E.g., color-red is MI only during phase 2, and is MC primarily only during Phase 3 (since phase 1 is so sparsely represented). In my experience, EEG noise is highly structured across a session and easily exploited by decoders. Plus, behavior seems quite different between Phase 2 and Phase 3. So, it seems likely that the classes you are asking the decoder to separate are highly confounded with temporally structured noise.

I suggest thinking of how to handle this concern in a rigorous way. A compelling way to address this would be to perform "cross-phase" decoding, however I am not sure if that is possible given the design.

Thank you for raising this important issue. To test whether decoding might be confounded by temporally structured noise, we performed a control decoding analysis. As the reviewer correctly pointed out, cross-phase decoding is not possible due to the experimental design. Alternatively, to maximize temporal separation between the training and test data, we divided the EEG data in phase 2 and phase 1&3 into the first and second half chronologically. Phase 1 and 3 were combined because they share the same MC and MI assignments. We then trained the decoders on one half and tested them on the other half. Finally, we averaged the decoding results across all possible assignments of training and test data. The similar patterns (Supplementary Fig.1) observed confirmed that the decoding results are unlikely to be driven by temporally structured noise in the EEG data. The clarification has been added to page 13 of the revised manuscript.

(1b) The time courses also seem concerning. What are we to make of the SR and SC timecourses, which have aggregate decoding dynamics that look to be <1Hz?

As detailed in the response to your next comment, some new results using data without baseline correction show a narrower time window of above-chance decoding. We speculate that the remaining results of long-lasting above-chance decoding could be attributed to trials with slow responses (some responses were made near the response deadline of 1500 ms). Additionally, as shown in Figure 6a, the long-lasting above-chance decoding seems to be driven by color and congruency representations. Thus, it is also possible that the binding of color and congruency contributes to decoding. This interpretation has been added to page 17 of the revised manuscript.

(1c) Some sanity checks would be one place to start. Time courses were baselined, but this is often not necessary with decoding; it can cause bias (10.1016/j.jneumeth.2021.109080), and can mask deeper issues. What do things look like when not baselined? Can variables be decoded when they should not be decoded? What does cross-temporal decoding look like - everything stable across all times, etc.?

As the reviewer mentioned, baseline-corrected data may introduce bias to the decoding results. Thus, we cited the van Driel et al (2021) paper in the revised manuscript to justify the use of EEG data without baseline-correction in decoding analysis (Page 27 of the revised manuscript), and re-ran all decoding analysis accordingly. The new results revealed largely similar results (Fig. 2, 4, 6 and 8 in the revised manuscript) with the following exceptions: narrower time window for separatable SC subspace and SR subspace (Fig. 4b), narrower time window for concurrent representations of SC and SR (Fig. 6a-b), and wider time window for the correlations of SC/SR representations with RTs (Fig. 8).

(2) The nature of the shared features between SR and SC subspaces is unclear.

The simulation is framed in terms of the amount of overlap, revealing the number of shared dimensions between subspaces. In reality, it seems like it's closer to 'proportion of volume shared', i.e., a small number of dominant dimensions could drive a large degree of alignment between subspaces.

What features drive the similarity? What features drive the distinctions between SR and SC? Aside from the temporal confounds I mentioned above, is it possible that some low-dimensional feature, like EEG congruency effect (e.g., low-D ERPs associated with conflict), or RT dynamics, drives discriminability among these classes? It seems plausible to me - all one would need is non-homogeneity in the size of the congruency effect across different items (subject-level idiosyncracies could contribute: 10.1016/j.neuroimage.2013.03.039).

Thank you for this question. To test what dimensions are shared between SC and SR subspaces, we first identify which factors can be shared across SC and SR subspaces. For SC, the eight conditions are the four colors × ISPC. Thus, the possible shared dimensions are color and ISPC. Additionally, because the four colors and words are divided into two groups (e.g., red-blue and green-yellow, counterbalanced across subjects, see Methods), the group is a third potential shared dimension. Similarly, for SR decoders, potential shared dimensions are word, ISPC and group. Note that each class in SC and SR decoders has both congruent and incongruent trials. Thus, congruency is not decodable from SC/SR decoders and hence unlikely to be a shared dimension in our analysis. To test the effect of sharing for each of the potential dimensions, we performed RSA on decoding results of the SC decoder trained on SR subspace (SR | SC) (Supplementary Fig. 4a) and the SR decoder trained on SC subspace (SC | SR) (Supplementary Fig. 4b), where the decoders indicated the decoding accuracy of shared SC and SR representations. In the SC classes of SR | SC, word red and blue were mixed within the same class, same were word yellow and green. The similarity matrix for “Group” of SR | SC (Supplementary Fig. 4a) shows the comparison between two word groups (red & blue vs. yellow & green). The similarity matrix for “Group” of SC | SR (Supplementary Fig. 4b) shows the comparison between two color groups (red & blue vs. yellow & green).

The RSA results revealed that the contributions of group to the SC decoder (Supplementary Fig. 5a) and the SR decoder (Supplementary Fig. 5b) were significant. Meanwhile, a wider time window showed significant effect of color on the SC decoder (approximately 100 - 1100 ms post-stimulus onset, Supplementary Fig. 5a) and a narrower time window showed significant effect of word on SR decoder (approximately 100 - 500 ms post-stimulus onset, Supplementary Fig. 5b). However, we found no significant effect of ISPC on either SC or SR decoders. We also performed the same analyses on response-locked data from the time window -800 to 200 ms. The results showed shared representation of color in the SC decoder (Supplementary Fig. 5c) and group in both decoders (Supplementary Fig. 5c-d). Overall, the above results demonstrated that color, word and group information are shared between SC and SR subspaces.

Lastly, we would like to stress that our main hypothesis for the cross-subspace decoding analysis is that SR and SC subspaces are not identical. This hypothesis was supported by lower decoding accuracy for cross-subspace than within-subspace decoders and enables following analyses that treated SC and SR as separate representations.

We have added the interpretation to page 13-14 of the revised manuscript.

(3) The time-resolved within-trial correlation of RSA betas is a cool idea, but I am concerned it is biased. Estimating correlations among different coefficients from the same GLM design matrix is, in general, biased, i.e., when the regressors are non-orthogonal. This bias comes from the expected covariance of the betas and is discussed in detail here (10.1371/journal.pcbi.1006299). In short, correlations could be inflated due to a combination of the design matrix and the structure of the noise. The most established solution, to cross-validate across different GLM estimations, is unfortunately not available here. I would suggest that the authors think of ways to handle this issue.

Thank you for raising this important issue. Because the bias comes from the covariance between the regressors and the same GLM was applied to all time points in our analysis, we assume that the inflation would be similar at different time points. Therefore, we calculated the correlation of SC and SR betas ranging from -200 to 0 ms relative to stimulus onset as a baseline (i.e., no SC or SR representation is expected before the stimulus onset) and compared the post-stimulus onset correlation coefficients against this baseline. We hypothesized that if the positively within-trial correlation of SC and SR betas resulted from the simultaneous representation instead of inflation, we should observe significantly higher correlation when compared with the baseline. To examine this hypothesis, we first performed the linear discriminant analysis (Supplementary Fig. 7a) and RSA regression (Supplementary Fig. 7b) on the -200 - 0 ms window relative to stimulus onset. We then calculated the average rbaseline of SC and SR betas on that time window for each participant (group results at each time point are shown in Supplementary Fig. 7c) and computed the relative correlation at each post-stimulus onset time point using (fisher-z (r) - fisher-z (rbaseline)). Finally, we performed a simple t test at the group level on baseline-corrected correlation coefficients with Bonferroni correction. The results (Fig. 6c) showed significantly more positive correlation from 100 - 500 ms post-stimulus onset compared with baseline, supporting our hypothesis that the positive within-trial correlation of SC and SR betas arise from simultaneous representation rather than inflation. The related interpretation was added to page 17 of the revised manuscript.

(4) Are results robust to running response-locked analyses? Especially the EEG-behavior correlation. Could this be driven by different RTs across trials & trial-types? I.e., at 400 ms poststim onset, some trials would be near or at RT/action execution, while others may not be nearly as close, and so EEG features would differ & "predict" RT.

Thanks for this question. We now pair each of the stimulus-locked EEG analysis in the manuscript with response-locked analysis. To control for RT variations among trial types, when using the linear mixed model (LMM) to predict RTs from trial-wise RSA results, we included a separate intercept for each of the eight trial types in SC or SR. Furthermore, at each time point, we only included trials that have not generated a response (for stimulus-locked analysis) or already started (for response-locked analysis). All the results (Fig. 3, 5, 7, 9 in the revised manuscript) are in support of our hypothesis. We added these detailed to page 31 of the revised manuscript.

(5) I suggest providing more explanation about the logic of the subspace decoding method - what trialtypes exactly constitute the different classes, why we would expect this method to capture something useful regarding ISPC, & what this something might be. I felt that the first paragraph of the results breezes by a lot of important logic.

In general, this paper does not seem to be written for readers who are unfamiliar with this particular topic area. If authors think this is undesirable, I would suggest altering the text.

To improve clarity, we revised the first paragraph of the SC and SR association subspace analysis to list the conditions for each of the SC and SR decoders and explain more about how the concept of being separatable can be tested by cross-decoding between SC and SR subspaces. The revised paragraph now reads:

“Prior to testing whether controlled and non-controlled associations were represented simultaneously, we first tested whether the two representations were separable in the EEG data.

In other words, we reorganized the 16 experimental conditions into 8 conditions for SC (4 colors × MC/MI, while collapsing across SR levels) and SR (4 words × 2 possible responses per word, while collapsing across SC levels) associations separately. If SC and SR associations are not separable, it follows that they encode the same information, such that both SC and SR associations can be represented in the same subspace (i.e., by the same information encoded in both associations). For example, because (1) the word can be determined by the color and congruency and (2) the most-likely response can be determined by color and ISPC, the SR association (i.e., association between word and most-likely response) can in theory be represented using the same information as the SC association. On the other hand, if SC and SR associations are separable, they are expected to be represented in different subspaces (i.e., the information used to encode the two associations is different). Notably, if some, but not all, information is shared between SC and SR associations, they are still separable by the unique information encoded. In this case, the SC and SR subspaces will partially overlap but still differ in some dimensions. To summarize, whether SC and SR associations are separable is operationalized as whether the associations are represented in the same subspace of EEG data. To test this, we leveraged the subspace created by the LDA (see Methods). Briefly, to capture the subspace that best distinguishes our experimental conditions, we trained SC and SR decoders using their respective aforementioned 8 experimental conditions. We then projected the EEG data onto the decoding weights of the LDA for each of the SC and SR decoders to obtain its respective subspace. We hypothesized that if SC and SR subspaces are identical (i.e., not separable), SC/SR decoding accuracy should not differ by which subspace (SC or SR) the decoder is trained on. For example, SC decoders trained in SC subspace should show similar decoding performance as SC decoders trained in SR subspace. On the other hand, if SC and SR association representations are in different subspaces, the SC/SR subspace will not encode all information for SR/SC associations. As a result, decoding accuracy should be higher using its own subspace (e.g., decoding SC using the SC subspace) than using the other subspace (e.g., decoding SC using the SR subspace). We used cross-validation to avoid artificially higher decoding accuracy for decoders using their own subspace (see Methods).” (Page 11-12).

We also explicitly tested what information is shared between SC and SR representations (see response to comment #2). Lastly, to help the readers navigate the EEG results, we added a section “Overview of EEG analysis” to summarize the EEG analysis and their relations in the following manner:

“EEG analysis overview. We started by validating that the 16 experimental conditions (8 unique stimuli × MC/MI) were represented in the EEG data. Evidence of representation was provided by above-chance decoding of the experimental conditions (Fig. 2-3). We then examined whether the SC and SR associations were separable (i.e., whether SC and SR associations were different representations of equivalent information). As our results supported separable representations of SC and SR association (Fig. 4-5), we further estimated the temporal dynamics of each representation within a trial using RSA. This analysis revealed that the temporal dynamics of SC and SR association representations overlapped (Fig. 6a-b, Fig. 7a-b). To explore the potential reason behind the temporal overlap of the two representations, we investigated whether SC and SR associations were represented simultaneously as part of the task representation, independently from each other, or competitively/exclusively (e.g., on some trials only SC association was represented, while on other trials only SR association was represented). This was done by assessing the correlation between the strength of SC and SR representations across trials (Fig. 6c, Fig. 7c). Lastly, we tested how SC and SR representations facilitated performance (Fig.8-9).” (Page 8-9).

Minor suggestions:

(6) I'd suggest using single-trial RSA beta coefficients, not t-values, as they can be more stable (it's a t-value based on 16 observations against 9 or so regressors.... the SE can be tiny).

Thank you for your suggestion. To choose between using betas and t-values, we calculate the proportion of outliers (defined as values beyond mean ± 5 SD) for each predictor of the design matrix and each subject. We found that outliers were less frequent for t-values than for beta coefficients (t-values: mean = 0.07%, SD = 0.009%; beta-values: mean = 0.19%, SD = 0.033%). Thus, we decided to stay with t-values.

(7) Instead of prewhitening the RTs before the HLM with drift terms, try putting those in the HLM itself, to avoid two-stage regression bias.

Thank you for your suggestion. Because our current LMM included each of the eight trial types in SC or SR as separate predictors with their own intercepts (as mentioned above), adding regressors of trial number and mini blocks (1-100 blocks) introduced collinearity (as ISPC flipped during the experiment). We therefore excluded these regressors from the current LMM (Page 31).

(8) The text says classical MDS was performed on decoding *accuracy* - is this accurate?

We now clarify in the manuscript that it is the decoders’ probabilistic classification results (Page 28).

(9) At a few points, it was claimed that a negative correlation between SC and SR would be expected within single trials, if the two were temporally dissociable. Wouldn't it also be possible that they are not correlated/orthogonal?

We agree with the reviewer and revised the null hypothesis in the cross-trial correlation analysis to include no correlation as SC and SR association representations may be independent from each other (Page 17, 22).

Reviewer #2 (Public review):

Summary:

In this EEG study, Huang et al. investigated the relative contribution of two accounts to the process of conflict control, namely the stimulus-control association (SC), which refers to the phenomenon that the ratio of congruent vs. incongruent trials affects the overall control demands, and the stimulus-response association (SR), stating that the frequency of stimulusresponse pairings can also impact the level of control. The authors extended the Stroop task with novel manipulation of item congruencies across blocks in order to test whether both types of information are encoded and related to behaviour. Using decoding and RSA, they showed that the SC and SR representations were concurrently present in voltage signals, and they also positively co-varied. In addition, the variability in both of their strengths was predictive of reaction time. In general, the experiment has a solid design, but there are some confounding factors in the analyses that should be addressed to provide strong support for the conclusions.

Strengths:

(1) The authors used an interesting task design that extended the classic Stroop paradigm and is potentially effective in teasing apart the relative contribution of the two different accounts regarding item-specific proportion congruency effect, provided that some confounds are addressed.

(2) Linking the strength of RSA scores with behavioural measures is critical to demonstrating the functional significance of the task representations in question.

Thank you for your positive feedback. We hope our responses below address your concerns.

Weakness:

(1) While the use of RSA to model the decoding strength vector is a fitting choice, looking at the RDMs in Figure 7, it seems that SC, SR, ISPC, and Identity matrices are all somewhat correlated. I wouldn't be surprised if some correlations would be quite high if they were reported. Total orthogonality is, of course, impossible depending on the hypothesis, but from experience, having highly covaried predictors in a regression can lead to unexpected results, such as artificially boosting the significance of one predictor in one direction, and the other one to the opposite direction. Perhaps some efforts to address how stable the timed-resolved RSA correlations for SC and SR are with and without the other highly correlated predictors will be valuable to raising confidence in the findings.

Thank you for this important point. The results of proportion of variability explained shown in the Author response table 1 below, indicated relatively higher correlation of SC/SR with Color and Identity. We agree that it is impossible to fully orthogonalize them. To address the issue of collinearity, we performed a control RSA by removing predictors highly correlated with others. Specifically, we calculated the variance inflation factor (VIF) for each predictor. The Identity predictor had a high VIF of 5 and was removed from the RSA. All other predictors had VIFs < 4 and were kept in the RSA. The results (Supplementary Fig. 6) showed patterns similar to the results with the Identity predictor, suggesting that the findings are not significantly influenced by collinearity. We have added the interpretation to page 17 of the revised manuscript.

(2) In "task overview", SR is defined as the word-response pair; however, in the Methods, lines 495-496, the definition changed to "the pairing between word and ISPC" which is in accordance with the values in the RDMs (e.g., mccbb and mcirb have similarity of 1, but they are linked to different responses, so should they not be considered different in terms of SR?). This needs clarification as they have very different implications for the task design and interpretation of results, e.g., how correlated the SC and SR manipulations were.

Thank you for pointing out this important issue with how our operationalization captures the concept in questions. In the revised manuscript, we clarified the stimulus-response (SR) association is the link between the word and the most-likely response (i.e., not necessarily the actual response on the current trial). This association is likely to be encoded based on statistical learning over several trials. On each trial, the association is updated based on the stimulus and the actual response. Over multiple trials, the accumulated association will be driven towards the most-common (i.e., most-likely) response. In our ISPC manipulation, a color is presented in mostly congruent/incongruent (MC/MI) trials, which will also pair a word with a most-likely response. For example, if the color blue is MC, the color blue, which leads to the response blue, will co-occur with the word blue with high frequency. In other words, the SR association here is between the word blue and the response blue. As the actual response is not part of the SR association, in the RDM two trial types with different responses may share the same SR association, as long as they share the same word and the same ISPC manipulation, which, by the logic above, will produce the same most-likely response. These clarifications have been added to page 4 and 29 of the revised manuscript.

In the revised manuscript (Page 17), we addressed how much the correlated SC and SR predictors in the RDM could affect the correlation analysis between SC and SR association representation strength. Specifically, we conducted the RSA using the same GLM on EEG data prior to stimulus onset (Supplementary Fig. 7a-b). As no SC and SR associations are expected to be present before stimulus onset, the correlation between SC and SR representation would serve as a baseline of inflation due to correlated predictors in the GLM (Supplementary Fig. 7c, also see comment #3 of R1). The SC-SR correlation coefficients following stimulus onset was then compared to the baseline to control for potential inflation (Fig. 6c). Significantly above-baseline correlation was still observed between ~100-500 ms post-stimulus onset, providing support for the hypothesis that SC and SR are encoded in the same task representation.

Minor suggestions:

(3) Overall, I find that calling SC-controlled and SR-uncontrolled representations unwarranted. How is the level controlledness defined? Both are essentially types of statistical expectation that provide contextual information for the block of tasks. Is one really more automatic and requires less conscious processing than the other? More background/justification could be provided if the authors would like to use these terms.

Following your advice, we have added more discussion on how controlledness is conceptualized in this work and in the literature, which reads:

“We consider SC and SR as controlled and uncontrolled respectively based on the literature investigating the mechanism of ISPC effect. The SC account posits that the ISPC effect results from conflict and involves conflict adaptation, which requires the regulation of attention or control (Bugg & Hutchison, 2013; Bugg et al., 2011; Schmidt, 2018; Schmidt & Besner, 2008). On the other hand, the SR account argues that ISPC effect does not require conflict adaptation but instead reflects contingency leaning. That is, the response can be directly retrieved from the association between the stimulus and the most-likely response without top-down regulation of attention or control. As more empirical evidence emerged, researchers advocating control view began to acknowledge the role of associative learning in cognitive control regarding the ISPC effect (Abrahamse et al., 2016). SC association has been thought to include both automatic that is fast and resource saving and controlled processes that is flexible and generalizable (Chiu, 2019). Overall, we do not intend to claim that SC is entirely controlled or SR is completely automatic. We use SC-controlled and SR-uncontrolled representations to align with the original theoretical motivation and to highlight the conceptual difference between SC and SR associations.” (Page 24-25)

(4) Figures 3c and d: the figures could benefit from more explanation of what they try to show to the readers. Also for 3d, the dimensions were aligned with color sets and congruencies, but word identities were not linearly separable, at least for the first 3 axes. Shouldn't one expect that words can be decoded in the SR subspace if word-response pairs were decodable (e.g., Figure 3b)?

Thank you for the insightful observation. We now clarified that Fig. 3c and d in the original manuscript (Fig. 4c and d in the current manuscript) aim to show how each of the 8 trial types in the SC and SR subspaces are represented. The MDS approach we used for visualization tries to preserve dissimilarity between trial types when projecting from data from a high dimensional to a low dimensional space. However, such projection may also make patterns linearly separatable in high dimensional space not linearly separatable in low dimensional space. For example, if the word blue has two points (-1, -1) and (1, 1) and the word red has two points (-1, 1) and (1, -1), they are not linearly separatable in the 2D space. Yet, if they are projected from a 3D space with coordinates of (-1, -1, -0.1), (1, 1, -0.1), (-1, 1, 0.1) and (1, -1, 0.1), the two words can be linearly separatable using the 3rd dimension. Thus, a better way to test whether word can be linearly separated in SR subspace is to perform RSA on the original high dimensional space. We performed the RSA with word (Supplementary Fig. 2) on the SR decoder trained on the SR subspace. Note that in Fig. 3c and d of the original script (Fig. 4c and d in the current manuscript) there are two pairs of words that are not linearly separable: red-blue and yellow-green. Thus, we specifically tested the separability within the two pairs using the one predictor for each pair, as shown in Supplementary Fig. 2. The results showed that within both word pairs individual words were presented above chance level (Supplementary Fig. 3). Considering that the decoders are linear, this finding indicates linear separability of the word pairs in the original SR subspace. The clarification has been added to page 13 (the end of the second paragraph) of the revised manuscript.

References

Abrahamse, E., Braem, S., Notebaert, W., & Verguts, T. (2016). Grounding cognitive control in associative learning. Psychological Bulletin, 142(7), 693-728.doi:10.1037/bul0000047.

Bugg, J. M., & Hutchison, K. A. (2013). Converging evidence for control of color-word Stroop interference at the item level. Journal of Experimental Psychology:Human Perception and Performance, 39(2), 433-449. doi:10.1037/a0029145.

Bugg, J. M., Jacoby, L. L., & Chanani, S. (2011). Why it is too early to lose control in accounts of item-specific proportion congruency effects. Journal of Experimental Psychology: Human Perception and Performance, 37(3), 844-859. doi:10.1037/a0019957.

Chiu, Y.-C. (2019). Automating adaptive control with item-specific learning. In Psychology of Learning and Motivation (Vol. 71, pp. 1-37).

Schmidt, J. R. (2018). Evidence against conflict monitoring and adaptation: An updated review. Psychonomic Bulletin & Review, 26(3), 753-771. doi:10.3758/s13423018-1520-z.

Schmidt, J. R., & Besner, D. (2008). The Stroop effect: Why proportion congruent has nothing to do with congruency and everything to do with contingency. Journal of Experimental Psychology: Learning, Memory, and Cognition, 34(3), 514-523. doi:10.1037/0278-7393.34.3.514.

Changes from the earlier version to the revised version

Removed in the revised version: Abstract

Abstract

Understanding how task knowledge is encoded neurally is crucial for uncovering the mechanisms underlying adaptive behavior. Here, we test the theory that all task information is integrated into a conjunctive task representation by investigating whether this representation simultaneously includes two types of associations that can guide behavior: stimulus-response (non-controlled) associations and stimulus-control (controlled) associations that inform how task focus should be adjusted to achieve goal-directed behavior. We extended the classic item-specific proportion congruency paradigm to dissociate the electroencephalographic (EEG) representations of controlled and non-controlled associations. Behavioral data replicated previous findings of association-driven adaptive behaviors. Decoding analyses of EEG data further showed that associations of controlled and non-controlled information were represented concurrently and differentially. Brain-behavioral analyses also showed that the strength of both associations was associated with faster responses. These findings support the idea that controlled and non-controlled associations are governed by an integrated task representation to guide adaptive behaviors simultaneously.

Introduction

Earlier version: Cognitive control coordinates our thoughts and actions with internal goals 1,2 via the abilities of sustained task focus when faced with distractions (stability) and flexible adjustments between task sets based on varying situations (flexibility). Both stability and flexibility can be implemented in a proactive 3,4 or reactive 5,6 manner 7-13. For task focus indicating cognitive stability, conflict paradigms have been employed to explore how the brain focuses on task-relevant information while avoiding the distraction from potent task-irrelevant information 14-16. For example, in the Stroop task 14, participants are shown a color word and asked to respond to the ink color of the word while ignoring the meaning of the word. The ink color can be consistent (congruent trials) or inconsistent (incongruent trials) with the meaning of the word, resulting in a conflict effect measured by worse performance on incongruent trials than on congruent trials. Some models of cognitive control postulate that the conflict effect reflects the involvement of cognitive control, which suppresses the habitual but inappropriate response to the word and/or boosts the novel and goal-directed response to the ink color to resolve the conflict on incongruent trials 3,17.

Revised version: Cognitive control coordinates our thoughts and actions with internal goals (Egner, 2017; Miller & Cohen, 2001) via the abilities of sustained task focus when faced with distractions (stability) and flexible adjustments between task sets based on varying situations (flexibility). Both stability and flexibility can be implemented in a proactive (Botvinick et al., 2001; Jiang et al., 2015) or reactive (Blais et al., 2007; Verguts & Notebaert, 2008) manner (Abrahamse et al., 2016; Braem & Egner, 2018; Braver, 2012; Brown et al., 2007; Egner, 2014, 2023; Jiang et al., 2018). For task focus indicating cognitive stability, conflict paradigms have been employed to explore how the brain focuses on task-relevant information while avoiding the distraction from potent task-irrelevant information (Eriksen & Eriksen, 1974; Simon & Rudell, 1967; Stroop, 1935). For example, in the Stroop task (Stroop, 1935), participants are shown a color word and asked to respond to the ink color of the word while ignoring the meaning of the word. The ink color can be consistent (congruent trials) or inconsistent (incongruent trials) with the meaning of the word, resulting in a conflict effect measured by worse performance on incongruent trials than on congruent trials. Some models of cognitive control postulate that the conflict effect reflects the involvement of cognitive control, which suppresses the habitual but inappropriate response to the word and/or boosts the novel and goal-directed response to the ink color to resolve the conflict on incongruent trials (Botvinick et al., 2001; Cohen et al., 1990).

Earlier version: A key question in the research on cognitive control is how the brain monitors control demands (i.e., how much control is needed). Recent research has demonstrated the importance of associative learning in guiding control demand 18-21. In particular, theories focusing on stimulus-control (SC) associations assume that the brain binds together a stimulus feature with the concurrent control demands 6,19,22-32. For example, in the Stroop task, if the color red is used on mostly congruent (MC) trials and the color blue is presented on mostly incongruent (MI) trials, then red would be associated with lower control demands relative to blue. This process has been proposed to account for the item-specific proportion congruency (ISPC) effect 33,34, whereby the conflict effect is reduced for stimulus features that predict higher control demands.

Revised version: A key question in the research on cognitive control is how the brain monitors control demands (i.e., how much control is needed). Recent research has demonstrated the importance of associative learning in guiding control demand (Bustamante et al., 2021; Jiang et al., 2020a; Jiang et al., 2020b; Wen et al., 2023). In particular, theories focusing on stimulus-control (SC) associations assume that the brain binds together a stimulus feature with the concurrent control demands (Bejjani et al., 2020; Blais et al., 2007; Bugg & Crump, 2012; Bugg & Hutchison, 2013; Bugg et al., 2011; Chiu et al., 2017; Colvett et al., 2025; Ileri-Tayar et al., 2024; Ileri-Tayar et al., 2022; Ileri-Tayar et al., 2025; Jiang et al., 2020a; Spinelli et al., 2022; Suh & Bugg, 2021). For example, in the Stroop task, if the color red is used on mostly congruent (MC) trials and the color blue is presented on mostly incongruent (MI) trials, then red would be associated with lower control demands relative to blue. This process has been proposed to account for the item-specific proportion congruency (ISPC) effect (Braem et al., 2019; Jacoby et al., 2003), whereby the conflict effect is reduced for stimulus features that predict higher control demands.

Earlier version: The ISPC effect is a natural prediction of stimulus-control learning, with the potential to provide a theoretical foundation for how cognitive control is scaffolded on structured representations of our environment. However, there have been theoretical challenges that the ISPC may arise from lower-level stimulus-response (SR) learning theory instead. This alternative account posits that associations may also form between stimulus and specific response in addition to more abstract SC associations 35-38. Returning to the Stroop task, if the word blue is often presented in red color, then the word blue will be associated with the ‘red’ response (in Stroop, often an assigned keypress or verbal response). Subsequent responses to this stimulus will be biased towards the associated response, whether or not they are correct.

Revised version: The ISPC effect is a natural prediction of stimulus-control learning, with the potential to provide a theoretical foundation for how cognitive control is scaffolded on structured representations of our environment. However, there have been theoretical challenges that the ISPC may arise from lower-level stimulus-response (SR) learning theory instead. This alternative account posits that associations may also form between stimulus and specific response (e.g., the most-likely response) in addition to more abstract SC associations (Atalay & Misirlisoy, 2012; Hazeltine & Mordkoff, 2014; Schmidt & Besner, 2008; Schmidt et al., 2016). Returning to the Stroop task, if the word blue is presented in red color (i.e., an incongruent trial), then the word blue will be associated with the ‘red’ response (in Stroop, often an assigned keypress or verbal response). Note that a stimulus may be associated with multiple responses. For example, if the word blue is presented in blue color (i.e., a congruent trial), then the word blue will also be associated with the ‘blue’ response. In this case, SR association is linked to the most-likely response. Following the example above, if the word blue is presented mostly on incongruent trials, then the word blue is paired more often with the red response than the blue response, and the SR is between the word blue and the red response. Subsequent responses to this stimulus will be biased towards the associated response, whether or not they are correct.

Earlier version: Since SC and SR accounts often predict similar behavioral patterns, previous studies have typically focused on one at a time through experimental controls, and current empirical evidence supports both SC 30,39-41 and SR 35 associations. Additionally, studies of ISPC have differentiated the behavioral effects and neural substrates of SC and SR associations when either SC or SR associations are available 23,29. For example, Chiu and colleagues 23 showed that the dynamics of SC and SR learning were captured by trial-wise prediction errors with a reinforcement learning model and that SC learning was encoded in the caudate nucleus while SR association was represented in the parietal cortex. Conceptually, SC associations affect behavior by modulating other cognitive processes via cognitive control, whereas SR associations drive behavior in a non-controlled manner by retrieving the associated response. These accounts provide competing explanations for what kind of knowledge is encoded in neural task representations. Task representations are often thought to be conjunctive, providing one-point access to all task-related information 5,42-51. However, it remains unexplored whether and how controlled and non-controlled information can be integrated into a task representation to guide adaptive behavior. If controlled and non-controlled information are encoded in the same representation, then: (1) SC and SR representations should occur simultaneously; (2) the strength of SC and SR representations, being in the same task presentation, should covary (i.e., positively correlated) over time as the strength of the task representation fluctuates and (3) both SC and SR associations should exhibit behavioral relevance.

Revised version: Since SC and SR accounts often predict similar behavioral patterns, previous studies have typically focused on one at a time through experimental controls, and current empirical evidence supports both SC (Bugg & Hutchison, 2013; Gonthier et al., 2016; Spinelli & Lupker, 2020a; Spinelli et al., 2019) and SR (Schmidt & Besner, 2008) associations. Additionally, studies of ISPC have differentiated the behavioral effects and neural substrates of SC and SR associations when either SC or SR associations are available (Bugg et al., 2011; Chiu et al., 2017). For example, Chiu and colleagues (Chiu et al., 2017) showed that the dynamics of SC and SR learning were captured by trial-wise prediction errors with a reinforcement learning model and that SC learning was encoded in the caudate nucleus while SR association was represented in the parietal cortex. Conceptually, SC associations affect behavior by modulating other cognitive processes via cognitive control, whereas SR associations drive behavior in a non-controlled manner by retrieving the associated response. These accounts provide competing explanations for what kind of knowledge is encoded in neural task representations. Task representations are often thought to be conjunctive, providing one-point access to all task-related information (Badre, 2024; Badre et al., 2021; Hommel, 2004, 2019; Kikumoto et al., 2024; Kikumoto & Mayr, 2020; Kikumoto et al., 2022a; Kikumoto et al., 2022b; Sakai, 2008; Schumacher & Hazeltine, 2016b; Verguts & Notebaert, 2008). However, it remains unexplored whether and how controlled and non-controlled information can be integrated into a task representation to guide adaptive behavior. If controlled and non-controlled information are encoded in the same representation, then: (1) SC and SR representations should occur simultaneously; (2) the strength of SC and SR representations, being in the same task presentation, should covary (i.e., positively correlated) over time as the strength of the task representation fluctuates and (3) both SC and SR associations should exhibit behavioral relevance.

Earlier version: To test these predictions, we extended the classic ISPC paradigm to dissociate each of SC and SR associations from their associated stimuli (e.g., separating SC association from the stimulus and the linked control demand). Using this new ISPC paradigm in an EEG experiment, we tested whether SC and SR form integrated task representations using decoding, representational subspace analysis, and representational similarity analysis (RSA). We found strong evidence supporting task representations that integrate controlled (SC) and non-controlled (SR) associations. First, time-resolved RSA on the decoding results showed simultaneous representations of SC and SR associations. Second, the strength of SC and SR associations was positively correlated, consistent with the hypothesis that both SC and SR associations are integrated in a task representation. Third, both the strength of SC and SR associations were associated with faster responses at the trial level, indicating the behavioral relevance of the neural representations. These results provide convergent evidence of an integrated task representation that contains both SC and SR associations to jointly guide goal-directed behaviors.

Revised version: To test these predictions, we extended the classic ISPC paradigm to dissociate each of SC and SR associations from their associated stimuli (e.g., separating SC association from the stimulus and the linked control demand). Using this new ISPC paradigm in an EEG experiment, we tested whether SC and SR form integrated task representations using decoding, representational subspace analysis, and representational similarity analysis (RSA). We found three main results on task representations that integrate controlled (SC) and non-controlled (SR) associations. First, time-resolved RSA on the decoding results showed simultaneous representations of SC and SR associations. Second, the strength of SC and SR associations was positively correlated, consistent with the hypothesis that both SC and SR associations are integrated in a task representation. Third, both the strength of SC and SR associations were associated with faster responses at the trial level, indicating the behavioral relevance of the neural representations. These results provide convergent evidence of an integrated task representation that contains both SC and SR associations to jointly guide goal-directed behaviors.

Results

Task overview

Earlier version: EEG data were acquired while participants (N = 40) performed a 4-key Stroop task. On each trial, participants responded to the color of a word while ignoring the identity of the word (e.g., RED presented in blue color; Fig. 1a). Participants were encouraged to respond as quickly and accurately as possible. We applied three key manipulations to this experiment (Fig. 1b). First, each stimulus could be either congruent or incongruent, defined by whether the color was consistent with the word. Second, to test whether controlled and non-controlled information are encoded in a task representation, we manipulated the ISPC 23,29,33,52 in Phase 1 to allow the associations between colors and their paired cognitive control demands (stimulus-control; SC) and the associations between words and their linked responses (stimulus-response; SR) to be induced at the same time. To keep the overall proportion of congruency at 50%, one set of colors was presented on 75% of congruent trials (mostly congruent trials, or MC), whereas the other set of colors was presented on 25% of congruent trials (mostly incongruent trials, or MI). The MC manipulation pairs color with low control state based on SC assumption (e.g., red-congruent) and word with consistent response (e.g., red-red) based on SR account. Similarly, the MI manipulation links color to high control state (e.g., blue-incongruent) and word to inconsistent response (e.g., yellow-green, because each set has two colors, the inconsistent response is unique). This standard ISPC manipulation can test whether neural representations of controlled and non-controlled information are wrapped on the same trial by combing with the following EEG analysis (See method). However, it potentially mixes color identity with SC, word identity with SR, and the ISPC between SC and SR. To de-confound these factors and obtain pure estimates of SC and SR on each trial, we modified this paradigm by flipping the ISPC contingencies across different phases of the task. For example, the color ‘red’ might be associated with 25% congruency in Phase 1, 75% congruency in Phase 2, and then return to 25% congruency in Phase 3 (see Methods). Overall, these manipulations resulted in a 2 (Congruency: congruent vs. incongruent) × 2 (ISPC: MC vs. MI) × 3 (Phase 1 - 3) within-subject factorial design.

Revised version: EEG data were acquired while participants (N = 40) performed a 4-key Stroop task. On each trial, participants responded to the color of a word while ignoring the identity of the word (e.g., RED presented in blue color; Fig. 1a). Participants were encouraged to respond as quickly and accurately as possible. We applied three key manipulations to this experiment (Fig. 1b). First, each stimulus could be either congruent or incongruent, defined by whether the color was consistent with the word. Second, to test whether controlled and non-controlled information are encoded in a task representation, we manipulated the ISPC (Bugg et al., 2011; Bustos et al., 2024; Chiu et al., 2017; Jacoby et al., 2003) in Phase 1 to allow the associations between colors and their paired cognitive control demands (stimulus-control; SC) and the associations between words and their most-likely response (i.e., not necessarily the actual response on the current trial; stimulus-response; SR) to be induced at the same time. To keep the overall proportion of congruency at 50%, one set of colors was presented on 75% of congruent trials (mostly congruent trials, or MC), whereas the other set of colors was presented on 25% of congruent trials (mostly incongruent trials, or MI). The MC manipulation pairs color with low control state based on SC assumption (e.g., red-congruent) and word with the consistent response (e.g., red-red) based on the SR account. Similarly, the MI manipulation links color to high control state (e.g., blue-incongruent) and word to the inconsistent response (e.g., yellow-green, because each set has two colors, the inconsistent response is unique). This standard ISPC manipulation can test whether neural representations of controlled and non-controlled information are wrapped on the same trial by combing with the following EEG analysis (See Methods). However, it potentially mixes color identity with SC, word identity with SR, and the ISPC between SC and SR. To de-confound these factors when estimating SC and SR association representations on each trial, we modified this paradigm by flipping the ISPC contingencies across different phases of the task. For example, the color ‘red’ might be associated with 25% congruency in Phase 1, 75% congruency in Phase 2, and then return to 25% congruency in Phase 3 (see Methods). Overall, these manipulations resulted in a 2 (Congruency: congruent vs. incongruent) × 2 (ISPC: MC vs. MI) × 3 (Phase 1 - 3) within-subject factorial design.

Experimental design and behavioral results (N = 40).

(a) Trial structure. (b) Procedure of the main experiment. The number following each stimulus indicates its trial count within a mini-block. (c) Group mean RT as a function of the ISPC effect. (d) Group mean error rate as a function of the ISPC effect. Error bars show standard errors of the mean (SEM). MI: mostly incongruent trials; MC: mostly congruent trials; incon: incongruent trials; con: congruent trials. *: p < 0.05; **: p < 0.01; ***: p < 0.001.

Behavioral patterns track manipulations of varying ISPC

RTs on correct trials within ±3 standard deviations (SD) of the mean and error rates were analyzed separately in each phase using a 2 (Congruency: congruent vs. incongruent) × 2 (ISPC: MC vs. MI) repeated-measures ANOVA. For RT data, there was a statistically significant interaction between Congruency and ISPC in each phase (Phase 1: F(1,39) = 36.36, p < 0.001, 𝜂𝜌2= 0.482; Phase 2: F(1,39) = 6.63, p = 0.014, 𝜂𝜌2 = 0.145; Phase 3: F(1,39) = 53.27, p < 0.001, 𝜂𝜌2 = 0.577), which replicated the ISPC effect of smaller congruency effect in the MI condition compared with that in the MC condition in each phase (Phase 1: t39 = 6.03, p < 0.001, Cohen’s d = 1.174; Phase 2: t39 = 2.58, p = 0.014, Cohen’s d = 0.476; Phase 3: t39 = 7.30, p < 0.001, Cohen’s d = 1.317; Fig. 1c). For error rate, the 2-way interaction was significant in Phase 1 (F(1,39) = 18.05, p < 0.001, 𝜂𝜌2 = 0.316) and Phase 3 (F(1,39) = 10.43, p < 0.01, 𝜂𝜌2= 0.211) and was driven by smaller conflict effects in the MI condition than in the MC condition (Phase 1: t39 = 4.25, p < 0.001, Cohen’s d = 0.842; Phase 3: t39 = 3.23, p < 0.01, Cohen’s d = 0.655; Fig. 1d), but was not significant in Phase 2 (F(1,39) = 1.50, p > 0.05). The above findings showed that the ISPC effect flipped from Phase 1 to Phase 2, and then reversed again in Phase 3, indicating that the experimental manipulations worked as expected. Additionally, we replicated the congruency effect (i.e., worse performance on incongruent than congruent trials) in both RT and error rate data in each phase (Supplementary Note 1 and Supplementary Table 1).

Added in the revised version: EEG analysis overview

Added in the revised version: We started by validating that the 16 experimental conditions (8 unique stimuli × MC/MI) were represented in the EEG data. Evidence of representation was provided by above-chance decoding of the experimental conditions (Fig. 2-3). We then examined whether the SC and SR associations were separable (i.e., whether SC and SR associations were different representations of equivalent information). As our results supported separable representations of SC and SR association (Fig. 4-5), we further estimated the temporal dynamics of each representation within a trial using RSA. This analysis revealed that the temporal dynamics of SC and SR association representations overlapped (Fig. 6a-b, Fig. 7a-b). To explore the potential reason behind the temporal overlap of the two representations, we investigated whether SC and SR associations were represented simultaneously as part of the task representation, independently from each other, or competitively/exclusively (e.g., on some trials only SC association was represented, while on other trials only SR association was represented). This was done by assessing the correlation between the strength of SC and SR representations across trials (Fig. 6c, Fig. 7c). Lastly, we tested how SC and SR representations facilitated performance (Fig. 8-9).

Removed in the revised version: Decodable and separate representations of controlled and non-controlled associations

Removed in the revised version: Before exploring the relationship between controlled and non-controlled representations, we first confirmed that each representation was encoded during task performance. We applied linear discriminant analysis (LDA) 53 with the features of 64-channel event-related potentials (ERP) data to decode the 16 experimental conditions (8 unique stimuli × MC/MI) at each time point. The decoding analysis was performed on a time window from stimulus onset to 1500ms post-stimulus onset (cf., Fig. 1a). We found that the decoding accuracy was significantly above chance level from ∼100ms after stimulus onset to the end of the analysis time window (Fig. 2a). To illustrate the distribution of data in the feature space, we applied multi-dimensional scaling (MDS, see Methods) to the ERP data averaged over 100 – 1500 ms after trial onset and visualized the centers of the 16 experimental conditions using the first 3 dimensions. The cumulative explained variance of the first 3 dimensions in the MDS was 0.61, 0.78 and 0.88, respectively. Dimensions 1 and 2 represented the 16 conditions in 8 SC classes (4 colors × MC/MI, Fig. 2b). MDS also categorized the conditions as 4 clusters denoting 4 colors along dimensions 1 and 3 (Fig. 2c).

Decoding performance including both SC and SR latent subspaces (N = 40).

(a) Group average decoding accuracy of all 16 experimental conditions as a function of time after stimulus onset. Shaded regions represent SEM. The solid line above the time axis denotes the time points showing above-chance (0.0625, or 1/16) decoding accuracy (p < 0.001, Bonferroni corrected). (b, c) MDS of EEG data across all experimental conditions. The label of each dot encodes the condition in the experimental design in the order of ISPC, congruency, color, and word. For example, “mccbr” means the condition with MC, congruent trial, blue color, and the word “red”. The color of each dot denotes the ink color of each condition.

Added in the revised version: Decoding performance including both SC and SR latent subspaces using response-locked analysis (N = 40).

Added in the revised version: (a) Group average decoding accuracy of all 16 experimental conditions as a function of time after response onset. Shaded regions represent SEM. The solid line above the time axis denotes the time points showing above-chance (0.0625, or 1/16) decoding accuracy (p < 0.001, Bonferroni corrected). (b, c) MDS of EEG data across all experimental conditions. The label of each dot encodes the condition in the experimental design in the order of ISPC, congruency, color, and word. For example, “mccbr” means the condition with MC, congruent trial, blue color, and the word “red”. The color of each dot denotes the ink color of each condition.

Removed in the revised version: Prior to testing whether controlled and non-controlled associations are represented simultaneously, we first tested whether the two representations are separable in the EEG data. To this end, we leveraged the subspace created by the LDA (see Methods). Briefly, to capture the subspace that best separates our experimental conditions, we projected our data onto the decoding weights of our LDA, forming 8-dimensional subspaces for SC (4 colors × MC/MI) and SR (4 words × 2 possible responses per word). We hypothesized that if SC and SR subspaces are identical, representations of SC and SR associations should not be separated. On the other hand, if SC and SR association representations are in different subspaces, the SC/SR subspace will not encode all information for SR/SC associations. As a result, decoding accuracy should be higher using its own subspace (e.g., decoding SC using the SC subspace) than using the other subspace (e.g., decoding SC using the SR subspace). We used cross-validation to avoid artificially higher decoding accuracy for decoders using their own subspace (see Methods). We verified this hypothesis in computational simulations to understand how the representational subspace analysis should perform under different accounts (see Methods). The simulation results (Fig. 3a) showed that: (1) when SC and SR subspaces are completely separable (i.e., the two subspaces shared no dimensions), decoding accuracy using the other subspace is at chance level; (2) when the SC and SR subspaces partially overlap (i.e., share some but not all dimensions), decoding accuracy using the other subspace is above chance level but less than when using its own subspace; and (3) when the SC and SR subspaces fully overlap (i.e., share all dimensions), decoding accuracy using the other subspace is above chance level and is comparable to that when using its own subspace.

Partially overlapping SC and SR subspaces.

(a) Simulation results (N = 40) of decoding accuracy as a function of the degree of subspace overlap, subspace type, and decoder. Both SC and SR subspaces are 8-dimensional (SC: 4 colors × MC/MI; SR: 4 words × 2 possible responses per word). The number of Shared dimensions indicates how many dimensions overlap between SC and SR subspaces. Each condition label is encoded in the format of “subspace | decoder”. For example, “SC | SR” means a SR decoder trained on the SC subspace. (b) Group average decoding accuracy over time as a function of which subspace the decoders are trained on. Shaded regions represent the SEM. Blue (orange) points denote the time points showing significantly better decoding accuracy when using the same subspace than when using the other subspace (p < 0.001, Bonferroni corrected). (c) MDS of the SC subspace. Each dot represents the center of a SC class. Dot color and label encode the ink color and cognitive control state, respectively. (d) MDS of the SR subspace. Each dot represents the center of a SR class. The label and dot color encode the word meaning and associated response, respectively.

Added in the revised version: Partially overlapping SC and SR subspaces using response-locked analysis.

Added in the revised version: (a) Group average decoding accuracy over time as a function of which subspace the decoders are trained on. Shaded regions represent the SEM. Blue (orange) points denote the time points showing significantly better decoding accuracy when using the same subspace than when using the other subspace (p < 0.001, Bonferroni corrected). (b) MDS of the SC subspace. Each dot represents the center of a SC class. Dot color and label encode the ink color and cognitive control state, respectively. (c) MDS of the SR. Each dot represents the center of a SR class. The label and dot color encode the word meaning and associated response, respectively.

Added in the revised version: Simultaneous EEG representations of SC and SR associations.

Added in the revised version: (a) Group average t values of representational strength for each factor over time. Squares below the lines indicate the significant time points (p < 0.001, Bonferroni corrected). (b) SC and SR association results from Fig. 6a. Shaded areas denote SEM. (c) Cross-trial correlation coefficient of representational strength between SC and SR associations, plotted as a function of time after stimulus onset. Shaded areas indicate SEM. Squares below the lines indicate the time points significantly above the baseline correlation strength (i.e., the average correlation coefficients in Supplementary Fig. 7c, p < 0.05, Bonferroni corrected).

Added in the revised version: Simultaneous EEG representations of SC and SR associations using response-locked analysis.

Added in the revised version: (a) Group average t values of representational strength for each factor over time. Squares below the lines indicate the significant time points (p < 0.001, Bonferroni corrected). (b) SC and SR association results from Fig. 7a. Shaded areas denote SEM. (c) Cross-trial correlation coefficient of representational strength between SC and SR associations, plotted as a function of time after response onset. Shaded areas indicate SEM. Squares below the lines indicate the time points significantly above the baseline correlation strength (i.e., the average correlation coefficients in Supplementary Fig. 7c, p < 0.05, Bonferroni corrected).

Added in the revised version: Both the strengths of SC and SR associations are correlated with RT.

Added in the revised version: (a) Group averaged t-values for each factor predicting RT in the LMM analysis. (b) SC and SR association results from Fig. 8a.

Added in the revised version: Correlations of SC and SR association strengths with RT using response-locked analysis.

Added in the revised version: (a) Group averaged t-values for each factor predicting RT in the LMM analysis. (b) SC and SR association results from Fig. 9a.

Added in the revised version: Decodable and separate representations of controlled and non-controlled associations

Added in the revised version: Before exploring the relationship between controlled and non-controlled representations, we first confirmed that each representation was encoded during task performance. We applied linear discriminant analysis (LDA) (James et al., 2013) with the features of 64-channel event-related potentials (ERP) data to decode the 16 experimental conditions (8 unique stimuli × MC/MI) at each time point. The decoding analysis was performed on a time window from stimulus onset to 1500 ms post-stimulus onset (cf., Fig. 1a). We found that the decoding accuracy was significantly above chance level from ∼100 to 1400 ms after stimulus onset (Fig. 2a). To illustrate the distribution of data in the feature space, we applied multi-dimensional scaling (MDS, see Methods) to the ERP data averaged over 100 - 1400 ms after trial onset and visualized the centers of the 16 experimental conditions using the first 3 dimensions. The cumulative explained variance of the first 3 dimensions in the MDS was 0.60, 0.76 and 0.87, respectively. Both the 8 SC classes of 16 conditions and the conditions as 4 color clusters were categorized along dimension 1 and dimension 2 (Fig. 2b) or along dimension 1 and dimension 3 (Fig. 2c). We also performed the above analyses on response-locked data and found similar patterns (Fig. 3): decoding accuracy was significantly above the chance from ∼ -600 to ∼200 ms relative to response (Fig. 3a); The cumulative explained variance of the first 3 dimensions in the MDS was 0.55, 0.71 and 0.85, respectively, revealing that MDS categorized the conditions as 4 clusters along dimension 1 and 2 (Fig. 3b) and separated 8 SC classes of 16 conditions (Fig. 3c).

Added in the revised version: Prior to testing whether controlled and non-controlled associations were represented simultaneously, we first tested whether the two representations were separable in the EEG data. In other words, we reorganized the 16 experimental conditions into 8 conditions for SC (4 colors × MC/MI, while collapsing across SR levels) and SR (4 words × 2 possible responses per word, while collapsing across SC levels) associations separately. If SC and SR associations are not separable, it follows that they encode the same information, such that both SC and SR associations can be represented in the same subspace (i.e., by the same information encoded in both associations). For example, because (1) the word can be determined by the color and congruency and (2) the most-likely response can be determined by color and ISPC, the SR association (i.e., association between word and most-likely response) can in theory be represented using the same information as the SC association. On the other hand, if SC and SR associations are separable, they are expected to be represented in different subspaces (i.e., the information used to encode the two associations is different). Notably, if some, but not all, information is shared between SC and SR associations, they are still separable by the unique information encoded. In this case, the SC and SR subspaces will partially overlap but still differ in some dimensions. To summarize, whether SC and SR associations are separable is operationalized as whether the associations are represented in the same subspace of EEG data. To test this, we leveraged the subspace created by the LDA (see Methods). Briefly, to capture the subspace that best distinguishes our experimental conditions, we trained SC and SR decoders using their respective aforementioned 8 experimental conditions. We then projected the EEG data onto the decoding weights of the LDA for each of the SC and SR decoders to obtain its respective subspace. We hypothesized that if SC and SR subspaces are identical (i.e., not separable), SC/SR decoding accuracy should not differ by which subspace (SC or SR) the decoder is trained on. For example, SC decoders trained in SC subspace should show similar decoding performance as SC decoders trained in SR subspace. On the other hand, if SC and SR association representations are in different subspaces, the SC/SR subspace will not encode all information for SR/SC associations. As a result, decoding accuracy should be higher using its own subspace (e.g., decoding SC using the SC subspace) than using the other subspace (e.g., decoding SC using the SR subspace). We used cross-validation to avoid artificially higher decoding accuracy for decoders using their own subspace (see Methods). We verified this hypothesis in computational simulations to understand how the representational subspace analysis should perform under different accounts (see Methods). The simulation results (Fig. 4a) showed that: (1) when SC and SR subspaces are completely separable (i.e., the two subspaces shared no dimensions), decoding accuracy using the other subspace is at chance level; (2) when the SC and SR subspaces partially overlap (i.e., share some but not all dimensions), decoding accuracy using the other subspace is above chance level but less than when using its own subspace; and (3) when the SC and SR subspaces fully overlap (i.e., share all dimensions), decoding accuracy using the other subspace is above chance level and is comparable to that when using its own subspace.

Earlier version: In the EEG data (Fig. 3b), we observed that the decoders using the other subspace (i.e., SC | SR and SR | SC) showed significantly above-chance accuracy (p < 0.001, Bonferroni corrected) after ∼100ms following stimulus onset. During the same time window, the decoder performance was significantly worse than the decoders using their own subspace (i.e., SC | SC and SR | SR). Thus, the empirical results are most consistent with the simulations of partially overlapped SC and SR spaces, suggesting that SC and SR association representations are separable in the EEG data. To visualize how data are represented in each subspace, we applied MDS to the SC (Fig. 3c, cumulative explained variance was 0.67 and 0.81 for the first 2 dimensions, respectively) and SR (Fig. 3d, cumulative explained variance was 0.75, 0.96, 0.99 for the first 3 dimensions) subspaces on EEG data averaged between 100 to 1500ms post-stimulus onset, which revealed the different structures of SC and SR encoded in the brain.

Revised version: In the EEG data (Fig. 4b; See Fig. 5a from response-locked analysis), we observed that the decoders using the other subspace (i.e., SC | SR and SR | SC) showed significantly above-chance accuracy (p < 0.001, Bonferroni corrected) after ∼100ms following stimulus onset. During the same time window, the decoder performance was significantly worse than the decoders using their own subspace (i.e., SC | SC and SR | SR). This pattern was not likely to be driven by temporally structured noise induced by different MC and MI assignments on different phases in the design (Supplementary Fig. 1). Thus, the empirical results are most consistent with the simulations of partially overlapped SC and SR spaces, suggesting that SC and SR association representations are separable in the EEG data. To show how each of the 8 conditions in the SC and SR subspaces are represented. The MDS approach was used for visualization to preserve dissimilarity between conditions when projecting from data from a high dimensional to a low dimensional space. The results showed that the 8 conditions were separate in both SC subspace (Fig. 4c, cumulative explained variance was 0.75, 0.88, 0.94 for the first 3 dimensions, respectively; See Fig. 5b from response-locked analysis, cumulative explained variance was 0.69, 0.86, 0.95 for the first 3 dimensions, respectively) and SR subspace (Fig. 4d, cumulative explained variance was 0.84, 0.93, 0.98 for the first 3 dimensions, respectively; See Fig. 5c from response-locked analysis, cumulative explained variance was 0.82, 0.95,0.99 for the first 3 dimension, respectively), which revealed the different structures of SC and SR encoded in the brain. We conducted RSA on SR decoders on the SR subspace and confirmed that these word pairs were indeed linearly separable in the SR subspace (Supplementary Fig. 2, 3).

Added in the revised version: To characterize the nature of the shared features between SR and SC, we first identified which factors can be shared across SC and SR subspaces. For SC, the eight conditions are the four colors × ISPC. Thus, the possible shared dimensions are color and ISPC. Additionally, because the four colors and words are divided into two groups (e.g., red-blue and green-yellow, counterbalanced across subjects, see Methods), the group is a third potential shared dimension. Similarly, for SR decoders, potential shared dimensions are word, ISPC and group. Note that each class in SC and SR decoders has both congruent and incongruent trials. Thus, congruency is not decodable from SC/SR decoders and hence unlikely to be a shared dimension in our analysis. To test the effect of sharing for each of the potential dimensions, we performed RSA on decoding results of the SC decoder trained on SR subspace (SR | SC) (Supplementary Fig. 4a) and the SR decoder trained on SC subspace (SC | SR) (Supplementary Fig. 4b). The RSA results revealed that the contributions of group to the SC decoder (Supplementary Fig. 5a) and the SR decoder (Supplementary Fig. 5b) were significant. Meanwhile, a wider time window showed significant effect of color on the SC decoder (approximately 100 - 1100 ms post-stimulus onset, Supplementary Fig. 5a) and a narrower time window showed significant effect of word on SR decoder (approximately 100 - 500 ms post-stimulus onset, Supplementary Fig. 5b). However, we found no significant effect of ISPC on either SC or SR decoders. We also performed the same analyses on response-locked data from the time window -800 to 200 ms. The results showed shared representation of color in the SC decoder (Supplementary Fig. 5c) and group in both decoders (Supplementary Fig. 5c-d). Overall, the above results demonstrated that color, word and group information are shared between SC and SR subspaces.

Added in the revised version: In summary, our main hypothesis for the cross-subspace decoding analysis is that SR and SC subspaces are not identical. This hypothesis was supported by lower decoding accuracy for cross-subspace than within-subspace decoders and enables following analyses that treated SC and SR as separate representations.

Simultaneously represented and positively correlated controlled and non-controlled associations

Earlier version: To further test whether SC and SR associations are simultaneously represented and contribute to adaptive task focus, we performed time-resolved RSA on decoding results to estimate the strength of representations of SC and SR associations while controlling for potential confounds (see Methods). In particular, since both SC and SR associations depend on the same ISPC condition, we included ISPC as a covariate in the same models as SC and SR associations to account for shared variability in the regression. The results showed that all the factors showed statistically significant effects on representational patterns during at least part of the trial time course (Fig. 4a). Crucially, both SC and SR effects remained statistically significant after ∼200ms following the stimulus onset (Fig. 4a-b), suggesting that SC and SR associations were simultaneously represented.

Revised version: To further test whether SC and SR associations are simultaneously represented and contribute to adaptive task focus, we performed time-resolved RSA on decoding results to estimate the strength of representations of SC and SR associations while controlling for potential confounds (see Methods). In particular, since both SC and SR associations depend on the same ISPC condition, we included ISPC as a covariate in the same models as SC and SR associations to account for shared variability in the regression. The results showed that all the factors showed statistically significant effects on representational patterns (Fig. 6a). Notably, statistically significant representations of color and congruency were observed for about 1s. This long period of representation may be attributed to the trials with slow responses, as both color and congruency are crucial task information to produce the correct response and to suppress distraction from the word. Another non-exclusive possibility is that color and congruency information remained present for post-trial learning of the SC association. Crucially, both SC and SR effects remained statistically significant from ∼250 to ∼ 650 ms following the stimulus onset (Fig. 6a-b). To control for collinearity between the predictors in the RSA (see Fig. 11), we calculated the variance inflation factor (VIF) for each predictor and repeated the same analysis using predictors with VIF scores no greater than 5 (only the Identity predictor was removed). We observed similar patterns (Supplementary Fig. 6). Similar patterns with response-locked analysis showed in Fig. 7a-b), suggesting that SC and SR associations were simultaneously represented.

Removed in the revised version: Simultaneous EEG representations of SC and SR associations.

Removed in the revised version: (a) Group average t values of representational strength for each factor over time. Squares below the lines indicate the significant time points (p < 0.001, Bonferroni corrected). (b) SC and SR association results from Fig. 5a. Shaded areas denote SEM. (c) Cross-trial correlation coefficient of representational strength between SC and SR associations, plotted as a function of time after stimulus onset. Shaded areas indicate SEM.

Earlier version: The integrated task representation hypothesis predicted that the representational strength of SC and SR associations should be positively correlated across trials. On the other hand, it remained possible that at the trial level SC and SR were not represented simultaneously. That is, a trial might only show an SR effect but no SC effect, or vice versa. If this were true, SC and SR effects should be negatively correlated across trials. We observed a statistically significant positive correlation over all time points (Fig. 4c), indicating that trials showing a strong SC effect tended to show stronger SR effects as well. This finding further supports the hypothesis of simultaneously encoded SC and SR associations within an integrated task representation.

Revised version: The integrated task representation hypothesis predicted that the representational strength of SC and SR associations should be positively correlated across trials. Alternatively, it remained possible that at the trial level SC and SR were not represented simultaneously. That is, a trial might only show an SR effect but no SC effect, or vice versa. If this were true, SC and SR effects should be negatively correlated across trials. A third possibility is that SC and SR association representations are independent from each other, leading to no correlation. To control for potential inflation of correlating regression coefficients from the same linear model (Cai et al., 2019), we computed correlation coefficients using pre-stimulus onset data as a baseline (Supplementary Fig. 7), as no SC or SR associations are expected to be present before the stimulus. We observed significantly above-baseline positive correlation between ∼100 - 500 ms following stimulus onset (Fig. 6c; See Fig. 7c from response-locked analysis). These findings indicated that trials showing a strong SC effect tended to show stronger SR effects as well, supporting the hypothesis of simultaneously encoded SC and SR associations within an integrated task representation.

Controlled and non-controlled associations jointly predict behavioral performance

Earlier version: To test whether the SC and SR effects from the RSA exhibited behavioral relevance, we predicted RT using a linear mixed model (LMM, see Methods) that included the representational strength of SC and SR associations, as quantified by the t-statistic from a trial-wise RSA. This analysis was conducted at each time point (time-locked to stimulus onset) prior to the group average RT (700.56ms) to test the effect of SC and SR representations on behavioral guidance. We found that faster responses were correlated with stronger representations of SC in the time window 380 – 520 ms (peak time: b = −0.009, SE = 0.002, t56581.22 = −4.93, p= 0.823-6) and SR in the time window 400 - 480 ms (peak time: b = −0.006, SE = 0.002, t56580.80 = −3.26, p= 0.001, Fig. 5a-b). These results suggest that controlled and non-controlled associations jointly influence performance before decision-making.

Revised version: To test whether the SC and SR effects from the RSA exhibited behavioral relevance, we predicted RT using a linear mixed model (LMM, see Methods) that included the representational strength of SC and SR associations, as quantified by the t-statistic from a trial-wise RSA. This analysis was conducted at each time point (time-locked to stimulus onset) prior to the group average RT (700.56ms) to test the effect of SC and SR representations on behavioral guidance. We found that faster responses were correlated with stronger representations of SC in the time window 380 - 700 ms (peak time: b = -0.046, SE = 0.007, t = -7.01) and SR in the time window 440 - 520 ms (peak time: b = -0.020, SE = 0.007, t = -2.93, Fig. 8a-b; See Fig. 9a-b for response-locked analysis; See Supplementary Fig. 8 and Supplementary Fig. 9 for the effects on frequent and infrequent trials). These results suggest that controlled and non-controlled associations jointly influence performance before decision-making.

Removed in the revised version: Both the strengths of SC and SR associations are correlated with RT.

Removed in the revised version: (a) Group averaged t-values for each factor predicting RT in the LMM analysis. (b) SC and SR association results from Fig. 5a.

Discussion

Adaptive behavior can be guided by associations for controlled processing (e.g., stimulus-control associations) and non-controlled processing (e.g., stimulus-response associations). To explore the relationship between these representations, we combined (1) a novel extension of a classic assay of associative learning in cognitive control with (2) neural subspace analyses of the time-resolved task representations. We found consistent support for distinct neural representations of SR and SC associations, but these associations appeared to share a synchronized temporal profile consistent with them constituting a common task representation.

Earlier version: We found behavioral ISPC effect in each phase of our task, even after the ISPC manipulations were flipped. This suggests that participants adjusted to the local ISPC conditions, which critically allowed us to dissociate stimulus features from their SC/SR associations. The behavioral ISPC pattern was consistent with previous studies 23-26,28-30,33. However, as both SC and SR associations are characterized by the same pattern of the ISPC effect, behavioral analysis and conventional contrasts between experimental conditions are insufficient to dissociate SC from SR effects. We addressed this issue with decoding, representational subspace analysis, and RSA on decoding results.

Revised version: We found behavioral ISPC effect in each phase of our task, even after the ISPC manipulations were flipped. This suggests that participants adjusted to the local ISPC conditions, which critically allowed us to dissociate stimulus features from their SC/SR associations. The behavioral ISPC pattern was consistent with previous studies (Bejjani et al., 2020; Bugg & Hutchison, 2013; Bugg et al., 2011; Chiu et al., 2017; Ileri-Tayar et al., 2024; Ileri-Tayar et al., 2022; Jacoby et al., 2003; Spinelli et al., 2022). However, as both SC and SR associations are characterized by the same pattern of the ISPC effect, behavioral analysis and conventional contrasts between experimental conditions are insufficient to dissociate SC from SR effects. We addressed this issue with decoding, representational subspace analysis, and RSA on decoding results.

Earlier version: We also found that SC and SR associations were encoded in distinct representations but appeared to be temporally synchronized as a task representation. An SR association itself is automatic behavior, but it is sometimes included as a core component of cognitive control 47. The original motivation for the ISPC task, and the SC theory, was to explore the role of associative learning in cognitive control beyond automatic S-R relationships 29,35,54. Brain-behavior analysis prior to the mean RT supported the prediction that SC and SR can jointly guide goal-directed behaviors. Previous fMRI studies with typical ISPC paradigms showed that the ISPC effect was related to cognitive control area including the dorsolateral prefrontal cortex (dlPFC), the dorsal anterior cingulate cortex (dACC) and the parietal cortex 55,56. Therefore, the results showing that SC and SR had distinct subspaces but emerged in the same time window should not suggest a paradox of findings or functional equivalence of SC and SR; rather, they highlight the brain’s ability to encode controlled and non-control information in parallel with partially overlapping representational subspaces to serve goal-directed behaviors. Additionally, the caudate nucleus associated with SC and the parietal cortex related with SR were observed in separate SC or SR conditions 23. It is possible that SC and SR are supported by partially overlapping neural substrates with similar temporal dynamics. Direct evidence is needed to validate the possibility.

Revised version: We also found that SC and SR associations were encoded in distinct representations but appeared to be temporally synchronized as a task representation. Brain-behavior analysis prior to the mean RT supported the prediction that SC and SR can jointly guide goal-directed behaviors. Previous fMRI studies with typical ISPC paradigms showed that the ISPC effect was related to cognitive control area including the dorsolateral prefrontal cortex (dlPFC), the dorsal anterior cingulate cortex (dACC) and the parietal cortex (Blais & Bunge, 2010; Grandjean et al., 2013). Therefore, the results showing that SC and SR had distinct subspaces but emerged in the same time window should not suggest a paradox of findings or functional equivalence of SC and SR; rather, they highlight the brain’s ability to encode controlled and non-control information in parallel with partially overlapping representational subspaces to serve goal-directed behaviors. Additionally, the caudate nucleus associated with SC and the parietal cortex related with SR were observed in separate SC or SR conditions (Chiu et al., 2017). It is possible that SC and SR are supported by partially overlapping neural substrates with similar temporal dynamics. Direct evidence is needed to validate the possibility.

Earlier version: We further explored how controlled and non-controlled associations jointly affect performance via brain-behavioral analyses with LMM. In general, we observed that the strength of SC and SR association representations are both negatively correlated with trial-level variability of RT. There could be one alternative explanation that either SC or SR, instead of both, occurs on each trial. If this were true, the coefficients of SC and SR from RSA results should be negatively correlated across trials (i.e., some trials show strong SC effects and weak SR effects, while other trials show the opposite). However, we found the opposite pattern, which indicates that there is no competition between SC and SR to be represented. One possibility leading to this positive correlation is that both associations are encapsulated in an integrated task representation 43,44,57. Specifically, if the task representation is strongly activated on a trial, both SC and SR associations will also be strongly triggered, leading to a positive correlation between SC and SR effects across trials.

Revised version: We further explored how controlled and non-controlled associations jointly affect performance via brain-behavioral analyses with LMM. In general, we observed that the strength of SC and SR association representations are both negatively correlated with trial-level variability of RT. There could be one alternative explanation that either SC or SR, instead of both, occurs on each trial. If this were true, the coefficients of SC and SR from RSA results should be negatively correlated or not correlated across trials (i.e., some trials show strong SC effects and weak SR effects, while other trials show the opposite). However, we found the opposite pattern, which indicates that there is no competition between SC and SR to be represented. One possibility leading to this positive correlation is that both associations are encapsulated in an integrated task representation (Hommel, 2004, 2019; Schumacher & Hazeltine, 2016a). Specifically, if the task representation is strongly activated on a trial, both SC and SR associations will also be strongly triggered, leading to a positive correlation between SC and SR effects across trials.

Earlier version: To interpret the ISPC effect, Bugg 58 proposed the dual item-specific mechanisms whereby SC and SR associations learning dominate the ISPC effect under different situations, which was supported by behavioral and imaging evidence 23,29,30. This means that SC associations would take over the adaptive effect when the frequency of incongruent trials is manipulated on task-relevant dimension, but SR associations would work when it is manipulated on task-irrelevant dimension. The results of the current study provide new temporal evidence for both control-driven ISPC and SR-driven ISPC and suggest that they can be represented simultaneously. Zooming in on the area of adaptive cognitive control, it has been debated whether adaptive task focus is actually involved with control or not, extending from ISPC to context/list-level PC 40,59-77 and the congruency sequence effect 78-84. Researchers have gradually formed an associative learning view on adaptive task focus 10,12,13. Egner 12 assumed integrating different levels of abstraction about associations including SR and SC into the same learning scheme serves for adaptive behaviors. The current study supports this assumption directly with the finding of joint action of control and non-control within the same stimulus for adaptive behaviors.

Revised version: To interpret the ISPC effect, Bugg (2015) proposed the dual item-specific mechanisms whereby SC and SR associations learning dominate the ISPC effect under different situations, which was supported by behavioral and imaging evidence (Bugg & Hutchison, 2013; Bugg et al., 2011; Chiu et al., 2017). This means that SC associations would take over the adaptive effect when the frequency of incongruent trials is manipulated on task-relevant dimension, but SR associations would work when it is manipulated on task-irrelevant dimension. The results of the current study provide new temporal evidence for both control-driven ISPC and SR-driven ISPC and suggest that they can be represented simultaneously. Zooming in on the area of adaptive cognitive control, it has been debated whether adaptive task focus is actually involved with control or not, extending from ISPC to context/list-level PC (Brosowsky & Crump, 2016, 2021; Bugg, 2014; Bugg & Chanani, 2011; Bugg et al., 2022; Cohen-Shikora et al., 2019; Crump, 2016; Crump et al., 2017; Crump et al., 2006; Crump & Milliken, 2009; Crump et al., 2008; Crump et al., 2018; Hutchison, 2011; Schmidt, 2013b; Schmidt, 2014b; Schmidt, 2017; Schmidt et al., 2014; Spinelli & Lupker, 2021; Spinelli et al., 2019; Weidler et al., 2022) and the congruency sequence effect (Brosowsky & Crump, 2018; Schmidt, 2013a; Schmidt, 2014a; Schmidt, 2018; Schmidt & Weissman, 2016; Spinelli & Lupker, 2020b; Weissman et al., 2016). Researchers have gradually formed an associative learning view on adaptive task focus (Abrahamse et al., 2016; Egner, 2014, 2023). Egner (2014) assumed integrating different levels of abstraction about associations including SR and SC into the same learning scheme serves for adaptive behaviors. The current study supports this assumption directly with the finding of joint action of control and non-control within the same stimulus for adaptive behaviors.

Earlier version: To better understand the role of complex associative learning in cognitive control, we developed a new way to test how representations are integrated, and how multiple representations of task information collaborate to guide goal-directed behaviors47,85. Previous research has employed partially overlapping tasks to test the conjunctive nature. For example, when there are two tasks producing different responses to the same input stimulus, encountering the stimulus will retrieve both task representations, leading to interference between the incompatible responses and impaired performance 86-89. However, this test is confounded by the retrieval of episodic memory (e.g., the last experiences performing the tasks, rather than the task representations), which is also conjunctive 43,44,90. Additional evidence supporting conjunctive task representation comes from studies showing that task identity is still decodable from electroencephalographic (EEG) data when accounting for individual task features such as stimulus, rule, and response 48-51,91. This approach relies on shared representations of task features across multiple tasks. It remained understudied how multiple task components within the same task are influenced by the conjunctive task representation. In this study, we extend the test of conjunctive task representation in two ways. First, we examine stochastic task features that must be learned from multiple episodes of task execution to tease apart retrieval of task representation from episodic retrieval. Second, we test the prediction that conjunctive task representation will simultaneously encode multiple task features instead of locating task-unique representations.

Revised version: To better understand the role of complex associative learning in cognitive control, we developed a new way to test how representations are integrated, and how multiple representations of task information collaborate to guide goal-directed behaviors(Badre, 2024; Musslick & Cohen, 2021). Previous research has employed partially overlapping tasks to test the conjunctive nature. For example, when there are two tasks producing different responses to the same input stimulus, encountering the stimulus will retrieve both task representations, leading to interference between the incompatible responses and impaired performance (Hubner & Druey, 2008; Koch et al., 2011; Korb et al., 2017; Rogers & Monsell, 1995). However, this test is confounded by the retrieval of episodic memory (e.g., the last experiences performing the tasks, rather than the task representations), which is also conjunctive (Altmann, 2011; Hommel, 2004, 2019). Additional evidence supporting conjunctive task representation comes from studies showing that task identity is still decodable from electroencephalographic (EEG) data when accounting for individual task features such as stimulus, rule, and response (Kikumoto et al., 2024; Kikumoto & Mayr, 2020; Kikumoto et al., 2022a; Kikumoto et al., 2022b; Rangel et al., 2023). This approach relies on shared representations of task features across multiple tasks. It remained understudied how multiple task components within the same task are influenced by the conjunctive task representation. In this study, we extend the test of conjunctive task representation in two ways. First, we examine stochastic task features that must be learned from multiple episodes of task execution to tease apart retrieval of task representation from episodic retrieval. Second, we test the prediction that conjunctive task representation will simultaneously encode multiple task features instead of locating task-unique representations.

Despite the theoretical differences between SR and SC associations, our subspace analysis revealed a pattern that is consistent with partially overlapping subspaces for SC and SR representations. SC and SR are different in their concepts and information processing, but they shared variance in their operational definitions by both depending on the same ISPC conditions. While this may have resulted in the partially overlapping representations, it may also reflect simultaneous retrieval to guide behavior. Critically, our RSA analyses revealed concurrent representations of SC and SR associations even when controlling for ISPC confounds. The experimental design combined with the RSA method ensures that the predictors for SC and SR associations are different.

Added in the revised version: An SR association itself is automatic behavior, but it is sometimes included as a core component of cognitive control (Badre, 2024). We consider SC and SR as controlled and uncontrolled respectively based on the literature investigating the mechanism of ISPC effect. The SC account posits that the ISPC effect results from conflict and involves conflict adaptation, which requires the regulation of attention or control (Bugg & Hutchison, 2013; Bugg et al., 2011; Schmidt, 2018; Schmidt & Besner, 2008). On the other hand, the SR account argues that ISPC effect does not require conflict adaptation but instead reflects contingency leaning. That is, the response can be directly retrieved from the association between the stimulus and the most-likely response without top-down regulation of attention or control. As more empirical evidence emerged, researchers advocating control view began to acknowledge the role of associative learning in cognitive control regarding the ISPC effect (Abrahamse et al., 2016). SC association has been thought to include both automatic that is fast and resource saving and controlled processes that is flexible and generalizable (Chiu, 2019). Overall, we do not intend to claim that SC is entirely controlled or SR is completely automatic. We use SC-controlled and SR-uncontrolled representations to align with the original theoretical motivation and to highlight the conceptual difference between SC and SR associations.

In summary, we find evidence of neural representations of associations of both controlled and non-controlled information quantified by SC and SR. The neural representations are represented simultaneously in an integrated representation in partially overlapping neural subspaces shortly after stimulus onset and remain active after the response is made. Moreover, the strengths of the associations prior to the response can jointly predict performance. These results shed light on how an integrated task representation, which consists of multiple associations that can drive behavior, can retrieve multiple associations of controlled and non-controlled information to guide adaptive behaviors.

Added in the revised version: Methods

Added in the revised version: Participants

Added in the revised version: Sixty-one healthy adults participated in the EEG experiment. Sixteen participants with excessive muscle artifacts (6 subjects), too many corrupted EEG channels (5 subjects), and missed partial data (5 subjects), and five participants with accuracy below 75% were removed, resulting in a final sample of 40 participants (28 females, 12 males; age mean = 23.9, SD = 6.2, age information was not acquired for the first 25 participants). All participants had normal or corrected-to-normal vision. All participants gave informed consent prior to the experiment, and either were paid or received course credits for participation. All procedures were approved by the University of Iowa Institutional Review Board (UIIRB #202001345).

Added in the revised version: Stimuli

Added in the revised version: The experiment was programmed using PsychoPy (ver. 2022.2.5). The stimulus was a color word presented in the center of the screen (24 inch LCD display with resolution of 1920× 1080) on a black background. Four color words (red, blue, yellow, and green) and their corresponding colors were used. For each participant, the colors and their corresponding words were divided into two sets of red-blue and yellow-green. Within each set, the words and colors formed a 2 × 2 factorial design, resulting in eight unique Stroop stimuli (Stroop, 1935) in total (e.g., red in red or blue color, blue in red or blue color, yellow in yellow or green color and green in yellow or green color).

Added in the revised version: Design and procedures

Added in the revised version: Each trial started with the presentation of a fixation cross for 500 ms, followed by a color word at the center of the screen (specified in ’height’ with 0.08 units where the full screen height equaled 1.0 unit) for 1500 ms or until a response was made (Fig. 1a). The four colors were mapped onto four keys (“H” for red, “J” for blue, “K” for yellow and “L” for green). Trials were separated by a black screen for the remainder of the 1500 ms response deadline plus a jittered (400/500/600 ms with equal probability) interval. The trials were grouped into mini-blocks of 16 trials each. The participants first underwent a practice session with 3 mini-blocks of 50% congruent trials. During practice, participants received feedback for a correct, incorrect or too slow response on each trial. After reaching the 75% accuracy criteria across all blocks, they moved on to the main experiment (Fig. 1b) with 8 mini-blocks for phase 1, 50 mini-blocks for phase 2 and 48 mini-blocks for phase 3.

Added in the revised version: Behavioral analysis

Added in the revised version: Behavioral data from the main experiment were analyzed. For response time (RT) data, only correct trials within ±3 SD from each subject’s mean RT were included in the analysis. Accuracy and RT were analyzed separately for each phase using a 2 (congruency) × 2 (ISPC) repeated-measures ANOVA.

Added in the revised version: EEG data acquisition and preprocessing

Added in the revised version: EEG data were recorded from a 64-channel active electrode cap with the ground at Fz and the reference at Pz using an actiCHamp amplifier (Brain Products) at a sampling rate of 500Hz with a hardware band-pass filter from 0.016 Hz to 1000 Hz. All electrode impedances werekept below 10 kΩ during the experiment. All EEG data were analyzed using the EEGLAB toolbox (https://sccn.ucsd.edu/eeglab/index.php) and custom MATLAB scripts. Preprocessing started with rejection and interpolation of noisy electrodes. Next, a 0.1–50 Hz band-pass filter was applied to the raw EEG data. The data were then re-referenced to the average of all electrodes and segmented into epochs spanning from -500 ms to 1500 ms relative to the stimulus onset. We used independent component analysis (ICA) to remove eye-blink and muscle artifacts from the epochs. These epochs were further segmented into smaller epochs ranging from -200 to 1500 ms relative to stimulus onset. These smaller epochs were also used to perform the response-locked analysis ranging from -800 to 200 ms relative to response. Epochs with amplitudes exceeding ± 80 μV were rejected. Fewer than 70% of all the epochs were excluded for each subject.

Added in the revised version: Decoding analyses

Added in the revised version: Decoding analysis was conducted on epoched EEG data without baseline correction to avoid potential bias in the decoding results (van Driel et al., 2021). To generate the dependent variable for the RSA, we first decoded the representations of all conditions using LDA (scikit-learn version 1.2.2 in Python (Pedregosa et al., 2011)). The preprocessed 64-channel ERP data were down-sampled to 50 Hz and were used as the features to decode the 16 experimental conditions (8 unique stimuli × MC/MI. Note that Phases1 and 3 shared the same ISPC manipulations). Decoding results were obtained by performing a 4-fold cross-validation procedure (Fig. 10). To avoid the bias of unbalanced trial numbers among different classes on decoding results, the more frequent trial types were down-sampled to be consistent with the less frequent trial types. Specifically, given the 75%/25% proportion congruency used in the manipulation of ISPC, the more frequent trial types had three times as many trials as the less frequent ones. Thus, the more frequent trial types were randomly divided into three parts, and one part was used for the 4-fold cross-validation. The cross-validation was repeated 30 times (i.e., 10 random partitions of the frequent trial types and each part being used once within each random partition). The LDA decoding accuracy was averaged. The decoding analysis was performed at each time point of a trial for each participant.

Added in the revised version: Illustration of linear discriminatory analysis (LDA).

Added in the revised version: To construct SC and SR subspaces, we trained 8-way SC (4 colors × MC/MI) and SR (4 words × 2 possible responses per word) LDA decoders and then transformed all training and test data into the subspaces of the LDA decoders. For each subspace, we built and tested SC and SR decoders using the same cross-validation approach as above. To avoid bias, all decoders used the same partition of training and test data, with the class label of each trial determined by its experimental condition. To test the statistical significance of decoding results, Bonferroni corrections with a family-wise error rate of α = 0.001 were applied to address the issue of multiple comparisons. Additionally, classical multidimensional scaling (MDS) analysis (Wickelmaier, 2003) on a 16×16 matrix of probabilistic classification results from the 16-way decoders was applied to visualize the neural representations of all conditions potentially including both SC and SR associations. Similar MDS analyses were performed on the 8×8 matrix of probabilistic classification results from the SC decoder and the other 8×8 matrix of decoding accuracy from the SR decoder to visualize their neural representations.

Added in the revised version: Simulations of cross-decoding accuracy

Added in the revised version: To simulate SC and SR subspaces with different levels of overlap, we first constructed non-overlapping SC and SR subspaces, each having 8 dimensions and consisting of 8 classes of 100 data points each. The centers of the classes were equidistant, with the distance set to the mean between-center distance in the EEG subspace. For each class, the data points were randomly drawn from a multivariate Gaussian distribution using the class center as the mean and the within-class standard deviation from the EEG subspace (averaged across all time points). Both SC and SR decoding were then conducted on each of the SC and SR subspaces using 4-fold cross-validation. To examine how overlap in subspaces affects decoding accuracy, the overlap between the two subspaces was manipulated by the number of shared dimensions. That is, if n (n ≤ 8) dimensions were shared, n dimensions of data were randomly chosen for SC_SR (i.e., SR decoder trained on the SC subspace) and SR_SC decoders. For SC_SC and SR_SR decoders, all 8 dimensions were used regardless of n.

Added in the revised version: RSA on decoding results

Added in the revised version: To further track the representational dynamics of SC and SR associations, we tested the hypothetical representational similarity patterns using the empirical representational similarity patterns from EEG data and linear regression-based RSA (Kikumoto & Mayr, 2020; Kikumoto et al., 2022a; Rangel et al., 2023). For each time point, its decoding results reflected the neural similarity between conditions (i.e., if two conditions have similar neural representations, the misclassification rate would increase) and were organized as the logit-transformed probability of a trial belonging to each of the 16 conditions to be the dependent variable. Six factors were considered, namely color, word, congruency, ISPC (MC/MI), SC (i.e., the pairing between color and ISPC), and SR (i.e., the pairing between word and ISPC), resulting in six hypothetical representational similarity patterns in the form of binary linear regressors (Fig. 11). Specifically, for each trial, each regressor was a column-vector of 16 cells, with each cell containing a binary value indicating whether each of the 16 conditions shared the feature (e.g., color) with the condition to which the current trial belonged to. Note that, as the SR association binds the word with the most-likely response rather than the actual response on the current trial, two trial types with different responses in the RDM (Fig. 11) may share the same SR association, as long as they share the same word and the same ISPC manipulation, which will produce the same most-likely response. Additionally, two nuisance binary regressors marking the condition and its frequency (i.e., whether a MC color on a congruent trial or a MI color on an incongruent trial) of the current trial were included. For each participant, the within-trial regressors formed a design matrix to regress against the empirical decoding results, resulting in trial-wise t values for each regressor. For each time point, individual mean t values for each regressor were submitted to a group-level one-sample t-test against 0. To test the significance of temporal representational strength, Bonferroni corrections with a family-wise error rate of α = 0.001 were applied to address the issue of multiple comparisons. Additionally, we correlated those regression coefficients of SC with that of SR across trials at each time point to test the possibility that SC and SR were not simultaneously represented within a trial.

Added in the revised version: Illustration of RSA.

Added in the revised version: The label of each row/column represents the condition in the experimental design including ISPC, congruency, color and word. For example, “mccbb” means the condition with MC, congruent trial, color blue and word blue. For each cell in a matrix, the color indicates whether the row and column conditions share the same factor (yellow = yes, blue = no) encoded by the matrix. For example, the cell at the 4th row and the 2nd column in the “SC” matrix encodes that the 4th condition (i.e., mcibr) and the 2nd condition (i.e., mccbb) share the same SC association.

Added in the revised version: Linear mixed model (LMM)

Added in the revised version: We further performed LMM at each time point preceding the mean RT across all trials (700.56 ms) to test whether trial-wise RSA results could predict the RT variability (Kikumoto & Mayr, 2020; Rangel et al., 2023). The model predicted trial-wise RT using each of the eight trial types in SC or SR (t-values from the RSA regressors) as separate predictors with their own intercepts. Log-transformed RTs were used as the dependent variable. At each time point, we only included trials that have not generated a response (for stimulus-locked analysis) or already started (for response-locked analysis).